***Important*** DO NOT CLEAR THE OUTPUT OF THIS NOTEBOOK AFTER EXECUTION!!!

```python
# if the following command generates an error, you probably didn't enable 
# the cluster security option "Allow API access to all Google Cloud services"
# under Manage Security → Project Access when setting up the cluster
!gcloud dataproc clusters list --region us-central1
```

Output:
```
NAME          PLATFORM  WORKER_COUNT  PREEMPTIBLE_WORKER_COUNT  STATUS   ZONE           SCHEDULED_DELETE
cluster-1937  GCE       4                                       RUNNING  us-central1-a
```

# Imports & Setup

```python
!pip install -q google-cloud-storage==1.43.0
!pip install -q graphframes
```

Output:
```
[33mWARNING: Running pip as the 'root' user can result in broken permissions and conflicting behaviour with the system package manager. It is recommended to use a virtual environment instead: https://pip.pypa.io/warnings/venv[0m[33m
[0m[33mWARNING: Running pip as the 'root' user can result in broken permissions and conflicting behaviour with the system package manager. It is recommended to use a virtual environment instead: https://pip.pypa.io/warnings/venv[0m[33m
[0m
```

```python
!touch log.txt
import logging
logging.basicConfig(level=logging.DEBUG,
                    filename='log.txt',
                    filemode='w',
                    format='%(asctime)s - %(name)s - %(levelname)s - %(message)s',
                    force=True)
logger = logging.getLogger(__name__)
```

```python
import pyspark
import sys
from collections import Counter, OrderedDict, defaultdict
import itertools
from itertools import islice, count, groupby
import pandas as pd
import os
import re
from operator import itemgetter
import nltk
from nltk.stem.porter import *
from nltk.corpus import stopwords
from time import time
from pathlib import Path
import pickle
import pandas as pd
from google.cloud import storage
from timeit import timeit
import numpy as np
import json
import hashlib
def _hash(s):
    return hashlib.blake2b(bytes(s, encoding='utf8'), digest_size=5).hexdigest()

nltk.download('stopwords')
```

Output:
```
[nltk_data] Downloading package stopwords to /root/nltk_data...
[nltk_data]   Unzipping corpora/stopwords.zip.
True
```

```python
# if nothing prints here you forgot to include the initialization script when starting the cluster
!ls -l /usr/lib/spark/jars/graph*
```

Output:
```
-rw-r--r-- 1 root root 247882 Jan  1 15:43 /usr/lib/spark/jars/graphframes-0.8.2-spark3.1-s_2.12.jar
```

```python

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[10], line 1
----> 1 from factory import *
      2 from Invertedindex import *

ModuleNotFoundError: No module named 'factory'
```

```python
from pyspark.sql import *
from pyspark.sql.functions import *
from pyspark import SparkContext, SparkConf, SparkFiles
from pyspark.sql import SQLContext
from graphframes import *
```

```python
spark
```

Output:
```
<pyspark.sql.session.SparkSession at 0x7f2b355bf550>
```

<!-- grade_id: cell-bucket_name, points: 0, grade: false, solution: true -->
```python
# Put your bucket name below and make sure you can access it without an error
bucket_name = '209100783-ir-project' 
full_path = f"gs://{bucket_name}/"
paths=[]

client = storage.Client()
blobs = client.list_blobs(bucket_name)
for b in blobs:
    if b.name != 'graphframes.sh':
        paths.append(full_path+b.name)
```

***GCP setup is complete!*** If you got here without any errors you've earned 10 out of the 35 points of this part.

# Building an inverted index

Here, we read the entire corpus to an rdd, directly from Google Storage Bucket and use your code from Colab to construct an inverted index.

```python
parquetFile = spark.read.parquet(*paths)

```

```python
doc_text_pairs = parquetFile.select("text", "id").rdd
```

We will count the number of pages to make sure we are looking at the entire corpus. The number of pages should be more than 6M

```python
# Count number of wiki pages
parquetFile.count()
```

Output:
```
6348910
```

Let's import the inverted index module. Note that you need to use the staff-provided version called `inverted_index_gcp.py`, which contains helper functions to writing and reading the posting files similar to the Colab version, but with writing done to a Google Cloud Storage bucket.

```python
# if nothing prints here you forgot to upload the file inverted_index_gcp.py to the home dir
%cd -q /home/dataproc
!ls 
```

Output:
```
Config.json  Invertedindex.py  assignment3_gcp.ipynb  factory.py
```

```python
# adding our python module to the cluster
sc.addFile("/home/dataproc/Config.json")
sc.addFile("/home/dataproc/Invertedindex.py")
sc.addFile("/home/dataproc/factory.py")

sys.path.insert(0,SparkFiles.getRootDirectory())
```

```python
from factory import *
from Invertedindex import *
```

```python
import json
with open("Config.json","r") as file_config:
    config_dict = json.load(file_config)
    index_config = config_dict['index_configuration']
    proj_config = config_dict['project_configuration']
    print(config_dict)
```

Output:
```
{'comments': {'type_stemmer_comment': 'the type of stemmer maybe [porter,snowball,...]', 'type_index_comment': 'the type of index maybe [text,anchor_text,title]', 'type_tokenizer_comment': 'the type of tokenizer maybe [python,spacy,regex(self),regex(web)]', 'type_similarity_function_comment': 'the type of similarity_function maybe [fastcosine,cosine (reut),spacy,numpy,....]', 'TF_MASK': '2 ** 16 - 1 # Masking the 16 low bits of an integer'}, 'index_configuration': {'BLOCK_SIZE': 1999998, 'TUPLE_SIZE': 6, 'TF_MASK': 65535, 'NUM_BUCKETS': 124, 'tokenizer': 'regex', 'stemmer': ['porter', 'snowball', 'regular'], 'type_of_index': ['text', 'title', 'anchor']}, 'project_configuration': {'project_id': 'ir-2022-my', 'bucket_name': '209100783-ir-project', 'bucket_to_upload': 'testing-bucketing', 'file_to_download': 'multistream10_preprocessed.parquet'}}
```

```python
Tok = FactoryIndex.get_tokenizer(index_config['tokenizer'])
english_stopwords = frozenset(stopwords.words('english'))
corpus_stopwords = ["category", "references", "also", "external", "links", 
                    "may", "first", "see", "history", "people", "one", "two", 
                    "part", "thumb", "including", "second", "following", 
                    "many", "however", "would", "became"]

all_stopwords = english_stopwords.union(corpus_stopwords)

NUM_BUCKETS = index_config['NUM_BUCKETS']

def token2bucket_id(token):
  return int(_hash(token),16) % NUM_BUCKETS
stemmering = None
global_doc_size = {}
```

```python
print(Tok)
```

Output:
```
<factory.TokenizerInterface object at 0x7f2ae84298e0>
```

<!-- grade_id: cell-token2bucket, points: 0, grade: false, solution: true -->
```python
# PLACE YOUR CODE HERE
def word_count(text, id,stemmer=None,idx_type=None):
  if stemmer is not None:
      logger.critical("this is stemmer")
      tokens = [stemmer.stem(token.group()) for token in Tok.tokenize(text.lower()) if token.group() not in all_stopwords]
  else:
      logger.critical("this is no stemmer")
      tokens = [token.group() for token in Tok.tokenize(text.lower()) if token.group() not in all_stopwords]

  token_countr = Counter(tokens)
  #calculating the vec size of the document
  if idx_type == 'text':
    global_doc_size[id] = np.linalg.norm(np.array(list(token_countr.values())))

  
  return [(word,(id,tf)) for word,tf in token_countr.items()]

def doc_to_count(text, id,stemmer=None,idx_type=None):
    if stemmer is not None:
        logger.critical("this is stemmer")
        tokens = [stemmer.stem(token.group()) for token in Tok.tokenize(text.lower()) if token.group() not in all_stopwords]
    else:
        logger.critical("this is no stemmer")
        tokens = [token.group() for token in Tok.tokenize(text.lower()) if token.group() not in all_stopwords]

    token_countr = Counter(tokens)
  
    doc_size = np.linalg.norm(np.array(list(token_countr.values())))
    doc_len = len(token_countr.values())

    
    return id,(np.round(doc_size,4),doc_len)

```

```python
def reduce_word_counts(unsorted_pl,thr_len_pl, idx_type):
 
#   if idx_type == 'text':
#     if len(unsorted_pl) < thr_len_pl:
#       logger.critical(f"list less then 50")
#       return []
      
  return sorted(unsorted_pl,key=lambda tup: tup[1],reverse=True)

```

```python
def partition_postings_and_write(postings,bucket_name,stem,index_type):

  rdd2 = postings.groupBy(lambda word: token2bucket_id(word[0])).map(lambda x : InvertedIndex.write_a_posting_list((x[0],x[1]),bucket_name,stem,index_type))
  return rdd2
```

```python
def preprocess_data(index_type,stem,bucket_name,path,index_name_file,w2df_file_name,doc_statistics,stemmering):
  client = storage.Client()
  t_start = time()
  if index_type == 'text':
    doc_to_len = doc_text_pairs.map(lambda x:doc_to_count(x[0],x[1],stemmering,index_type)).collectAsMap()
    avg_tot_len = __builtins__.sum(map(lambda x: x[1],doc_to_len.values()))/len(doc_to_len.keys())
    doc_to_len["AVG_DOC_LEN"] = avg_tot_len
    with open(doc_statistics,'wb') as f:
      pickle.dump(doc_to_len,f)
      bucket = client.bucket(bucket_name)
      blob_posting_locs = bucket.blob(f"postings_gcp/{path}{doc_statistics}")
      blob_posting_locs.upload_from_filename(f"{doc_statistics}")
      del doc_to_len

  word_counts = doc_text_pairs.flatMap(lambda x: word_count(x[0],x[1],stemmering,index_type))
  postings = word_counts.groupByKey().mapValues(lambda x :reduce_word_counts(x,0, index_type))


#   postings_filtered = postings.filter(lambda x: len(x[1])>0)

  w2df = postings.mapValues(lambda key: len(key))
  w2df_dict = w2df.collectAsMap()
  
  with open(w2df_file_name,'wb') as f:
      pickle.dump(w2df_dict,f)
      bucket = client.bucket(bucket_name)
      blob_posting_locs = bucket.blob(f"postings_gcp/{path}{w2df_file_name}")
      blob_posting_locs.upload_from_filename(f"{w2df_file_name}")
      del w2df_dict
  _ = partition_postings_and_write(postings_filtered,'testing-bucketing',stem,index_type).collect()
  print(time()-t_start)
  #-------------------------------------
  

```

```python
def create_inverted_index(path,index_name_file,w2df_file_name,doc_statistics):
  inverted = InvertedIndex()
  

  with open(w2df_file_name,'rb') as f:
    w2df_dict = pickle.load(f)
    inverted.df = w2df_dict
  with open(doc_statistics,'rb') as f:
    doc_stat = pickle.load(f)
    inverted.global_doc_size = doc_stat


  

  super_posting_locs = defaultdict(list)
  for blob in client.list_blobs('testing-bucketing', prefix=f'postings_gcp/{path}'):
    if not blob.name.endswith("pickle"):
      continue
    with blob.open("rb") as f:
      posting_locs = pickle.load(f)
      
      for k, v in posting_locs.items():
        # todo change to super_posting_locs[k] = v
        super_posting_locs[k].extend(v)

  inverted.posting_locs = super_posting_locs

  # write the global stats out
  inverted.write_index('.', index_name_file,'testing-bucketing',path)

```

```python

index_type = 'text'
stem = "regular"
print(stem)
if stem == 'regular':
    stemmering = None
else:
    stemmering = FactoryIndex.get_stemmer(stem)
client = storage.Client()
bucket_name = 'testing-bucketing'
path = f"{stem}-index/{index_type}-index/"
index_name_file = f'{stem}_{index_type}_index'
w2df_file_name = f"{stem}_{index_type}_w2df_statistics.stat"
doc_statistics = f"{stem}_{index_type}_doc_statistics.stat"
###############################################
preprocess_data(index_type,stem,bucket_name,path,index_name_file,w2df_file_name,doc_statistics,stemmering)
###############################################
create_inverted_index(path,index_name_file,w2df_file_name,doc_statistics)
###############################################



```

Output:
```
regular
this is creating the inverted index                                             
7667.524652957916
```

```python

index_type = 'text'
stem = "snowball"
print(stem)
if stem == 'regular':
    stemmering = None
else:
    stemmering = FactoryIndex.get_stemmer(stem)
client = storage.Client()
bucket_name = 'testing-bucketing'
path = f"{stem}-index/{index_type}-index/"
index_name_file = f'{stem}_{index_type}_index'
w2df_file_name = f"{stem}_{index_type}_w2df_statistics.stat"
doc_statistics = f"{stem}_{index_type}_doc_statistics.stat"
###############################################
preprocess_data(index_type,stem,bucket_name,path,index_name_file,w2df_file_name,doc_statistics,stemmering)
###############################################
create_inverted_index(path,index_name_file,w2df_file_name,doc_statistics)
###############################################




```

Output:
```
/home
dataproc
```

```python

index_type = 'text'
stem = "porter"
print(stem)
if stem == 'regular':
    stemmering = None
else:
    stemmering = FactoryIndex.get_stemmer(stem)
client = storage.Client()
bucket_name = 'testing-bucketing'
path = f"{stem}-index/{index_type}-index/"
index_name_file = f'{stem}_{index_type}_index'
w2df_file_name = f"{stem}_{index_type}_w2df_statistics.stat"
doc_statistics = f"{stem}_{index_type}_doc_statistics.stat"
###############################################
preprocess_data(index_type,stem,bucket_name,path,index_name_file,w2df_file_name,doc_statistics,stemmering)
###############################################
create_inverted_index(path,index_name_file,w2df_file_name,doc_statistics)
###############################################



```

```python
new_index = InvertedIndex().read_index('.',index_name_file)

```

Output:
```
this is creating the inverted index
```

```python
# for idx,i in enumerate(new_index.global_doc_size.items()):
#     print(i)
#     if idx >= 1000:
#         break

```

<!-- grade_id: cell-index_construction, points: 0, grade: false, solution: true -->
```python
# time the index creation time
t_start = time()
# word counts map
word_counts = doc_text_pairs.flatMap(lambda x: word_count(x[0], x[1]))
postings = word_counts.groupByKey().mapValues(reduce_word_counts)
# filtering postings and calculate df
postings_filtered = postings.filter(lambda x: len(x[1])>50)
w2df = calculate_df(postings_filtered)
w2df_dict = w2df.collectAsMap()
# partition posting lists and write out
_ = partition_postings_and_write(postings_filtered).collect()
index_const_time = time() - t_start
```

<!-- grade_id: cell-index_const_time, points: 10, grade: true, solution: false -->
```python
# test index construction time
assert index_const_time < 60*120
```

<!-- grade_id: collect-posting, points: 0, grade: true, solution: false -->
```python
# collect all posting lists locations into one super-set
super_posting_locs = defaultdict(list)
for blob in client.list_blobs(bucket_name, prefix='postings_gcp'):
  if not blob.name.endswith("pickle"):
    continue
  with blob.open("rb") as f:
    posting_locs = pickle.load(f)
    for k, v in posting_locs.items():
      super_posting_locs[k].extend(v)
```

Putting it all together

```python
# Create inverted index instance
inverted = InvertedIndex()
# Adding the posting locations dictionary to the inverted index
inverted.posting_locs = super_posting_locs
# Add the token - df dictionary to the inverted index
inverted.df = w2df_dict
# write the global stats out
inverted.write_index('.', 'index')
# upload to gs
index_src = "index.pkl"
index_dst = f'gs://{bucket_name}/postings_gcp/{index_src}'
!gsutil cp $index_src $index_dst
```

```python
!gsutil ls -lh $index_dst
```

# PageRank

**YOUR TASK (10 POINTS):** Compute PageRank for the entire English Wikipedia. Use your implementation for `generate_graph` function from Colab below.

```python
# Put your `generate_graph` function here
```

<!-- grade_id: cell-PageRank, points: 0, grade: false, solution: true -->
```python
t_start = time()
pages_links = spark.read.parquet("gs://wikidata_preprocessed/*").select("id", "anchor_text").rdd
# construct the graph 
edges, vertices = generate_graph(pages_links)
# compute PageRank
edgesDF = edges.toDF(['src', 'dst']).repartition(124, 'src')
verticesDF = vertices.toDF(['id']).repartition(124, 'id')
g = GraphFrame(verticesDF, edgesDF)
pr_results = g.pageRank(resetProbability=0.15, maxIter=6)
pr = pr_results.vertices.select("id", "pagerank")
pr = pr.sort(col('pagerank').desc())
pr.repartition(1).write.csv(f'gs://{bucket_name}/pr', compression="gzip")
pr_time = time() - t_start
pr.show()
```

<!-- grade_id: cell-PageRank_time, points: 10, grade: true, solution: false -->
```python
# test that PageRank computaion took less than 1 hour
assert pr_time < 60*60
```

# Reporting

**YOUR TASK (5 points):** execute and complete the following lines to complete 
the reporting requirements for assignment #3. 

```python
# size of input data
!gsutil du -sh "gs://wikidata_preprocessed/"
```

Output:
```
14.28 GiB    gs://wikidata_preprocessed
```

```python
# size of index data
index_dst = f'gs://{bucket_name}/postings_gcp/'
!gsutil du -sh "$index_dst"
```

Output:
```
5.93 GiB     gs://sise_ir_assignment3/postings_gcp
```

```python
# How many USD credits did you use in GCP during the course of this assignment?
cost = 0 
print(f'I used {cost} USD credit during the course of this assignment')
```

**Bonus (10 points)** if you implement PageRank in pure PySpark, i.e. without using the GraphFrames package, AND manage to complete 10 iterations of your algorithm on the entire English Wikipedia in less than an hour. 


<!-- grade_id: cell-PageRank_Bonus, points: 0, grade: false, solution: true -->
```python
#If you have decided to do the bonus task - please copy the code here 

bonus_flag = False # Turn flag on (True) if you have implemented this part

t_start = time()

# PLACE YOUR CODE HERE

pr_time_Bonus = time() - t_start

```

<!-- grade_id: cell-PageRank_Bonus-time, points: 10, grade: true, solution: false -->
```python
# Note:test that PageRank computaion took less than 1 hour
assert pr_time_Bonus < 60*60 and bonus_flag
```